# CS 3110/5990: Data Privacy
## Homework 4

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

```python
BASE_CLIP = 100 * 1000
```

## Question 1 (10 points)

Complete the definition of `dp_sum_capgain` below. Your definition should compute a differentially private sum of the "Capital Gain" column of the `adult` dataset, and have a total privacy cost of `epsilon`.

<!-- grade_id: cell-6a369f61d8eef943, points: 0, grade: false, solution: true -->
```python
def dp_sum_capgain(epsilon):
    clipped = adult['Capital Gain'].clip(upper=BASE_CLIP)
    return laplace_mech(clipped.sum(), sensitivity=BASE_CLIP, epsilon=epsilon)

dp_sum_capgain(1.0)
```

Output:
```
np.float64(35092796.990063794)
```

<!-- grade_id: cell-e5b1229ba2249f71, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1

real_sum = adult['Capital Gain'].sum()
r1 = np.mean([pct_error(real_sum, dp_sum_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_sum, dp_sum_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_sum, dp_sum_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)
print()
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 10
assert r2 < 2
assert r3 < 0.2
```

Output:
```
Average errors: 2.863941064097467 0.2881285309630241 0.029104264389763088
```

## Question 2 (10 points)

In 2-5 sentences each, answer the following:

- What clipping parameter did you use in your definition of `dp_sum_capital`, and why?

I used a clipping parameter of 100,000. I picked it by looking at the max (not always a great practice) and rounding it off.

- What was the sensitivity of the query you used in `dp_sum_capital`, and how is it bounded?

The sensitivity was 100,000. It is bounded by clipping each person to at most 100,000.

- Argue that your definition of `dp_sum_capital` has a total privacy cost of `epsilon`

The laplace mechanisim was applied once with a parameter of epsilon. 1 [applied 1 time] * 1 [epsilon of 1] = 1 [total privacy cost of 1]

<!-- grade_id: cell-986eec754f32c072, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (10 points)

Complete the definition of `dp_avg_capgain` below. Your definition should compute a differentially private average (mean) of the "Capital Gain" column of the adult dataset, and have a **total privacy cost of epsilon**.

<!-- grade_id: cell-93f2a5153dca0e95, points: 0, grade: false, solution: true -->
```python
def dp_avg_capgain(epsilon):
    pSum = dp_sum_capgain(epsilon/2)

    clipped = adult['Capital Gain'].clip(upper=BASE_CLIP)
    pCount = laplace_mech(len(clipped), sensitivity=1, epsilon=epsilon/2)

    avg = pSum / pCount

    return avg


dp_avg_capgain(1.0)
```

Output:
```
np.float64(1076.8883659791459)
```

<!-- grade_id: cell-ef77b325b5c58908, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3

real_avg = adult['Capital Gain'].mean()
r1 = np.mean([pct_error(real_avg, dp_avg_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_avg, dp_avg_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_avg, dp_avg_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 20
assert r2 < 4
assert r3 < 0.4
```

Output:
```
Average errors: 5.198239851418219 0.5893161048171935 0.059543348460570626
```

```python
def dp_avg_capgain_setClip(epsilon, clip):
    pSum = dp_sum_capgain(epsilon/2)

    clipped = adult['Capital Gain'].clip(upper=clip)
    pCount = laplace_mech(len(clipped), sensitivity=1, epsilon=epsilon/2)

    avg = pSum / pCount

    return avg

clips = [0, 50000, 75000, 99998, 100000]
for clip in clips:
    r1 = np.mean([pct_error(real_avg, dp_avg_capgain_setClip(0.1, clip)) for _ in range(100)])
    r2 = np.mean([pct_error(real_avg, dp_avg_capgain_setClip(1, clip)) for _ in range(100)])
    r3 = np.mean([pct_error(real_avg, dp_avg_capgain_setClip(10, clip)) for _ in range(100)])

    print(f"Average errors, cliping at ({clip}): ep = 0.1 {r1}, ep = 1 {r2}, ep = 10 {r3}")

print("")
print(f"All Adults: {len(adult)}")
for clip in clips:
    print(f"Capital Gain > {clip}: {len(adult[adult['Capital Gain'] > clip])}")
```

Output:
```
Average errors, cliping at (0): ep = 0.1 5.1972682478375365, ep = 1 0.48536238812379023, ep = 10 0.052984628196245165
Average errors, cliping at (50000): ep = 0.1 5.369325462041667, ep = 1 0.5019679767523991, ep = 10 0.05960465442334979
Average errors, cliping at (75000): ep = 0.1 5.56166160621419, ep = 1 0.5649422216353195, ep = 10 0.052276394243587765
Average errors, cliping at (99998): ep = 0.1 5.475926393477326, ep = 1 0.5571100800569617, ep = 10 0.05594141648952185
Average errors, cliping at (100000): ep = 0.1 5.25267341644209, ep = 1 0.5534018848046951, ep = 10 0.06925644951142697

All Adults: 32561
Capital Gain > 0: 2712
Capital Gain > 50000: 159
Capital Gain > 75000: 159
Capital Gain > 99998: 159
Capital Gain > 100000: 0
Please review the code in the cell(s) to identify a possible cause of the failure. 

Click <a href='https://aka.ms/vscodeJupyterKernelCrash'>here</a> for more info. 

View Jupyter <a href='command:jupyter.viewOutput'>log</a> for further details.
```

## Question 4 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your definition of `dp_avg_capgain` has a total privacy cost of `epsilon`

The laplace mechanisim was applied twice with a parameter of epsilon divided by 2. 2 [applied 2 times] * 0.5 [epsilon of 1/2] = 1 [total privacy cost of 1]

- For sums and averages, which seems to be more important for accuracy - the value of the clipping parameter $b$ or the scale of the noise added? Why?

Here the scale of the noise seems to make impact accuracy more than the value of the clipping parameter. 
- Do you think the answer to the previous point will be true for every dataset? Why or why not?

<!-- grade_id: cell-7ea16016366529bd, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 (20 points)

Write a function `auto_avg` that returns the differentially private average of a Pandas series `s`. Your function should automatically determine the clipping parameter `b`, and should enforce differential privacy for a **total privacy cost** of `epsilon`. You can assume that all values are non-negative (i.e. 0 or greater).

<!-- grade_id: cell-10b352b38a9b5f85, points: 0, grade: false, solution: true -->
```python
def auto_avg(data_set, epsilon):
    bs = [10 ** i for i in range(0, 7, 0.25)]
    epsilon_i = epsilon / len(bs)
    last_sum = 0
    for b in bs: #Satisfies len(bs) * epsilon_i-DP by sequential comp
        clipped_sum = data_set.clip(lower=0, upper = b).sum()
        dp_sum = laplace_mech(clipped_sum, sensitivity=b, epsilon=epsilon_i)
        print(dp_sum)
        if dp_sum < last_sum:
            return b
        else:
            last_sum = dp_sum
            
    raise Exception("No good val!")

```

<!-- grade_id: cell-49fb57b1ebfcc0da, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 5
r1 = np.mean([pct_error(adult['Age'].mean(), auto_avg(adult['Age'], 1.0)) for _ in range(20)])
r2 = np.mean([pct_error(adult['Capital Gain'].mean(), auto_avg(adult['Capital Gain'], 1.0)) for _ in range(20)])
r3 = np.mean([pct_error(adult['fnlwgt'].mean(), auto_avg(adult['fnlwgt'], 1.0)) for _ in range(20)])

print('Average errors:', r1, r2, r3)
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 1
assert r2 < 100
assert r3 < 1
```

Output:
```
Cell In[31], line 2, in <listcomp>(.0)
----> 2 # TEST CASE for question 5

Cell In[30], line 2, in auto_avg(data_set, epsilon)
      1 def auto_avg(data_set, epsilon):
----> 2     bs = [10 ** i for i in range(0, 7, 0.25)]
      3     epsilon_i = epsilon / len(bs)
      4     last_sum = 0
      5     for b in bs: #Satisfies len(bs) * epsilon_i-DP by sequential comp

TypeError: 'float' object cannot be interpreted as an integer
```

## Question 6

In 2-5 sentences each, answer the following:

- Explain your strategy for implementing `auto_avg`
- Argue informally that your definition of `auto_avg` has a total privacy cost of `epsilon`
- Did your solution work well for all three example columns? If it did not work well on any of them, why not?
- When is your solution likely to *not* work well? (i.e. what properties does the data have to have, in order for your solution to not work well?)

<!-- grade_id: cell-0ff63739ed66adf7, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE