# CS 3110/5110: Data Privacy
## Homework 2

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
adult = adult.dropna()
```

## Question 3 (10 points)

Write code to determine the `Education-Num` of the individual named Ardyce Golby by performing a differencing attack. Your code should *only* use aggregate data to find Ardyce's education number.

<!-- grade_id: cell-ef28e136429ae8c2, points: 0, grade: false, solution: true -->
```python
def ardyce_education():
    name = 'Ardyce Golby'
    column = 'Education-Num'
    adult_minus = adult[adult['Name'] != name]
    diff = adult[column].sum() - adult_minus[column].sum()
    return diff

```

<!-- grade_id: cell-d06fb722893a39a5, points: 10, grade: true, solution: false -->
```python
# TEST CASE for Question 3
assert ardyce_education() == 12
```

```python

```

## Question 1 (20 points)

Implement a more efficient version of `is_k_anonymous`. The inefficient implementation, taken from the textbook, appears below.

**Hint**: use the [`value_counts`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.value_counts.html) or `group_by` functions, and make sure no count is less than $k$.

```python
# Checking for k-Anonymity, taken from the textbook
# def is_k_anonymous(k, qis, df):
#     for index, row in df.iterrows():
#         query = ' & '.join([f'`{col}` == "{row[col]}"' for col in qis])
#         rows = df.query(query)
#         if (rows.shape[0] < k):
#             return False
#     return True
```

<!-- grade_id: cell-41bce010b027856e, points: 0, grade: false, solution: true -->
```python
# Checking for k-anonymity more efficiently
def is_k_anonymous(k, qis, df):
    """Returns true if df satisfies k-Anonymity for the quasi-identifiers 
    qis. Returns false otherwise."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3d5fd22faf4dd36c, points: 20, grade: true, solution: false -->
```python
# TEST CASES for question 1

assert not is_k_anonymous(2, ['Age'], adult)
assert is_k_anonymous(1, ['Age'], adult)
assert is_k_anonymous(1, ['Age', 'Occupation'], adult)
```

Output:
```
Cell In[58], line 3
      1 # TEST CASES for question 1
      2 
----> 3 assert not is_k_anonymous(2, ['Age'], adult)
      4 assert is_k_anonymous(1, ['Age'], adult)
      5 assert is_k_anonymous(1, ['Age', 'Occupation'], adult)

Cell In[57], line 6, in is_k_anonymous(k, qis, df)
      2 def is_k_anonymous(k, qis, df):
      3     """Returns true if df satisfies k-Anonymity for the quasi-identifiers 
      4     qis. Returns false otherwise."""
      5     # YOUR CODE HERE
----> 6     raise NotImplementedError()

NotImplementedError:
```

## Question 2 (10 points)

Consider the definition of `generalize` below, taken from the textbook. The function takes a dataframe `df` and a dictionary `depths` that describes how much to generalize each column of `df`. Generalizing a column to a depth of $n$ replaces the $n$ least-significant digits of each number in that column by zeroes. For example, we could generalize column `A` by making its least-significant digit a 0 and column `B` by doing the same for 2 digits with the following depth specification:

```python
depths = {
    'A': 1,
    'B': 2
}
```

```python
def generalize(df, depths):
    return df.apply(lambda x: x.apply(lambda y: int(int(y/(10**depths[x.name]))*(10**depths[x.name]))))
```

Using the `generalize` function, generalize the `Age` column of the `adult` dataset to a depth of 1. Drop the other columns of the dataset. Your result should achieve $k$-Anonymity for $k=20$.

<!-- grade_id: cell-ebd030727b205883, points: 0, grade: false, solution: true -->
```python
def generalize_adult_age():
    depths = {
        # YOUR CODE HERE
        raise NotImplementedError()
    }
    
    return generalize(adult[['Age']], depths)
```

<!-- grade_id: cell-040f45c8ba9921c4, points: 10, grade: true, solution: false -->
```python
assert is_k_anonymous(20, ['Age'], generalize_adult_age())
```

## Question 3 (10 points)

Using the `generalize` function, generalize the `Age` and `Zip` columns of the `adult` dataset in order to achieve $k$-Anonymity for $k=5$. Your result should drop other columns besides these two.

<!-- grade_id: cell-2a5ad51eb81f11c9, points: 0, grade: false, solution: true -->
```python
def generalize_adult_age_zip():
    depths = {
        # YOUR CODE HERE
        raise NotImplementedError()
    }

    return generalize(adult[['Age', 'Zip']], depths)
```

<!-- grade_id: cell-fd1c63463e2a9665, points: 10, grade: true, solution: false -->
```python
assert is_k_anonymous(5, ['Age', 'Zip'], generalize_adult_age_zip())
```

## Question 4 (30 points)

In 1-4 sentences each, answer the following:

1. How much generalization was required to achieve $k=5$ in question 3?
2. Does this level of generalization significantly impact the utility of the $k$-Anonymized data? Why or why not?
3. Why is generalizing the `adult` dataset so challenging? (**Hint**: consider outliers)
4. Is there another approach, in addition to our simple generalization method, that might work better?
5. What is a simple method for generalizing the `Occupation` column?

<!-- grade_id: cell-44d08233c2b2e94a, points: 30, grade: true, solution: true -->
YOUR ANSWER HERE