# CS 3110/5110: Data Privacy
## Homework 1

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
adult_pii = adult[['Name', 'DOB', 'SSN', 'Zip', 'Age']]
adult_deid = adult.drop(columns=['Name', 'SSN'])
```

## Question 1 (20 points)

Using the dataframes `adult_pii` and `adult_deid`, write code to conduct a linking attack to recover the names of as many individuals in `adult_deid` as possible. Your solution should be parameterized by the set of columns to use in the attack.

<!-- grade_id: cell-e6389ceace711c32, points: 0, grade: false, solution: true -->
```python
def linking_attack(cols):
    return pd.merge(adult_pii, adult_deid, on=cols, how='inner')

```

<!-- grade_id: cell-251b9f1c6f7b4630, points: 20, grade: true, solution: false -->
```python
# TEST CASES for Question 1

assert len(linking_attack(['Zip'])) == 43191
assert len(linking_attack(['Zip', 'DOB'])) == 32563
assert len(linking_attack(['Zip', 'Age'])) == 32755
```

## Question 2 (10 points)

How many individuals in this dataset are uniquely identified by their Zip code? How many are uniquely identified by their age?

Hint: note that the number of *unique ZIP codes* is **different** from the number of *individuals uniquely identified by ZIP code*.

Hint: you can use the `value_counts` method to count the number of occurences of each value in a series.

<!-- grade_id: cell-0e380ac68dd0c77e, points: 0, grade: false, solution: true -->
```python
def unique_zipcode():
    # YOUR CODE HERE
    raise NotImplementedError()

def unique_dob():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9b009427c61dbb2b, points: 10, grade: true, solution: false -->
```python
# TEST CASES (INCOMPLETE) for Question 2
# INCOMPLETE means that additional test cases will be used to check your answer during grading

assert unique_zipcode() > 0
assert unique_dob() > 0

```