# CS 3110/5990: Data Privacy
## Homework 4

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

```python
BASE_CLIP = 100 * 1000
```

## Question 1 (10 points)

Complete the definition of `dp_sum_capgain` below. Your definition should compute a differentially private sum of the "Capital Gain" column of the `adult` dataset, and have a total privacy cost of `epsilon`.

<!-- grade_id: cell-6a369f61d8eef943, points: 0, grade: false, solution: true -->
```python
def dp_sum_capgain(epsilon):
    clipped = adult['Capital Gain'].clip(upper=BASE_CLIP)
    return laplace_mech(clipped.sum(), sensitivity=BASE_CLIP, epsilon=epsilon)

dp_sum_capgain(1.0)
```

Output:
```
np.float64(35148313.212613955)
```

<!-- grade_id: cell-e5b1229ba2249f71, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1

real_sum = adult['Capital Gain'].sum()
r1 = np.mean([pct_error(real_sum, dp_sum_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_sum, dp_sum_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_sum, dp_sum_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)
print()
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 10
assert r2 < 2
assert r3 < 0.2
```

Output:
```
Average errors: 3.345149809341347 0.2881807489413895 0.027326746695814204
```

## Question 2 (10 points)

In 2-5 sentences each, answer the following:

- What clipping parameter did you use in your definition of `dp_sum_capital`, and why?

I used a clipping parameter of 100,000. I picked it by looking at the max (not always a great practice) and rounding it off.

- What was the sensitivity of the query you used in `dp_sum_capital`, and how is it bounded?

The sensitivity was 100,000. It is bounded by clipping each person to at most 100,000.

- Argue that your definition of `dp_sum_capital` has a total privacy cost of `epsilon`

The laplace mechanisim was applied once with a parameter of epsilon. 1 [applied 1 time] * 1 [epsilon of 1] = 1 [total privacy cost of 1]

<!-- grade_id: cell-986eec754f32c072, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (10 points)

Complete the definition of `dp_avg_capgain` below. Your definition should compute a differentially private average (mean) of the "Capital Gain" column of the adult dataset, and have a **total privacy cost of epsilon**.

<!-- grade_id: cell-93f2a5153dca0e95, points: 0, grade: false, solution: true -->
```python
def dp_avg_capgain(epsilon):
    pSum = dp_sum_capgain(epsilon/2)

    clipped = adult['Capital Gain'].clip(upper=BASE_CLIP)
    pCount = laplace_mech(len(clipped), sensitivity=1, epsilon=epsilon/2)

    avg = pSum / pCount

    return avg


dp_avg_capgain(1.0)
```

Output:
```
np.float64(1070.1307818045545)
```

<!-- grade_id: cell-ef77b325b5c58908, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3

real_avg = adult['Capital Gain'].mean()
r1 = np.mean([pct_error(real_avg, dp_avg_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_avg, dp_avg_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_avg, dp_avg_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 20
assert r2 < 4
assert r3 < 0.4
```

Output:
```
Average errors: 6.485197783624728 0.606367708717932 0.06214869863307619
```

```python
def dp_avg_capgain_setClip(epsilon, clip):
    pSum = dp_sum_capgain(epsilon/2)

    clipped = adult['Capital Gain'].clip(upper=clip)
    pCount = laplace_mech(len(clipped), sensitivity=1, epsilon=epsilon/2)

    avg = pSum / pCount

    return avg

clips = [0, 50000, 75000, 99998, 100000]
for clip in clips:
    r1 = np.mean([pct_error(real_avg, dp_avg_capgain_setClip(0.1, clip)) for _ in range(100)])
    r2 = np.mean([pct_error(real_avg, dp_avg_capgain_setClip(1, clip)) for _ in range(100)])
    r3 = np.mean([pct_error(real_avg, dp_avg_capgain_setClip(10, clip)) for _ in range(100)])

    print(f"Average errors, cliping at ({clip}): ep = 0.1 {r1}, ep = 1 {r2}, ep = 10 {r3}")

print("")
print(f"All Adults: {len(adult)}")
for clip in clips:
    print(f"Capital Gain > {clip}: {len(adult[adult['Capital Gain'] > clip])}")
```

Output:
```
Average errors, cliping at (0): ep = 0.1 4.907580687127629, ep = 1 0.6148122352259259, ep = 10 0.06996980435627508
Average errors, cliping at (50000): ep = 0.1 4.998715764701733, ep = 1 0.5515612701918993, ep = 10 0.06336831376893387
Average errors, cliping at (75000): ep = 0.1 5.899923306143168, ep = 1 0.5742632546330438, ep = 10 0.05376611700872583
Average errors, cliping at (99998): ep = 0.1 5.703001430054519, ep = 1 0.5742651567825532, ep = 10 0.06204562268889829
Average errors, cliping at (100000): ep = 0.1 5.849026742991031, ep = 1 0.5093742042959706, ep = 10 0.055738857544787086

All Adults: 32561
Capital Gain > 0: 2712
Capital Gain > 50000: 159
Capital Gain > 75000: 159
Capital Gain > 99998: 159
Capital Gain > 100000: 0
```

## Question 4 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your definition of `dp_avg_capgain` has a total privacy cost of `epsilon`

The laplace mechanisim was applied twice with a parameter of epsilon divided by 2. 2 [applied 2 times] * 0.5 [epsilon of 1/2] = 1 [total privacy cost of 1]

- For sums and averages, which seems to be more important for accuracy - the value of the clipping parameter $b$ or the scale of the noise added? Why?

Here the scale of the noise seems to make impact accuracy more than the value of the clipping parameter. 
- Do you think the answer to the previous point will be true for every dataset? Why or why not?

<!-- grade_id: cell-7ea16016366529bd, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 (20 points)

Write a function `auto_avg` that returns the differentially private average of a Pandas series `s`. Your function should automatically determine the clipping parameter `b`, and should enforce differential privacy for a **total privacy cost** of `epsilon`. You can assume that all values are non-negative (i.e. 0 or greater).

<!-- grade_id: cell-10b352b38a9b5f85, points: 0, grade: false, solution: true -->
```python
# Based on code from class
def pick_clip(data_set, epsilon):
    # I tested a view different bs lists, this one seemed to work the best.
    bs = range(1, 10 ** 4, 20)

    epsilon_i = epsilon / len(bs)
    last_sum = 0
    for b in bs: #Satisfies len(bs) * epsilon_i-DP by sequential comp
        clipped_sum = data_set.clip(lower=0, upper = b).sum()
        dp_sum = laplace_mech(clipped_sum, sensitivity=b, epsilon=epsilon_i)
        if dp_sum < last_sum:
            return b
        else:
            last_sum = dp_sum
    
    print("No good values! returning bs[-1]")
    return bs[-1]

def auto_avg(data_set, epsilon):
    # Divides epsilon into parts used by laplace_mech
    # I tested a view different ratios, this one seemed to work the best
    clip_budget = (1 * epsilon) / 4
    sum_budget = (3 * epsilon) / 8
    count_budget = (3 * epsilon) / 8
    # Total budget must add up to epsilon
    assert clip_budget + sum_budget + count_budget == epsilon


    clip_val = pick_clip(data_set, clip_budget)

    noisy_sum = laplace_mech(data_set.sum(), clip_val, sum_budget)
    noisy_count = laplace_mech(len(data_set), 1, count_budget)
    noisy_avg = noisy_sum / noisy_count

    return noisy_avg


```

<!-- grade_id: cell-49fb57b1ebfcc0da, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 5
r1 = np.mean([pct_error(adult['Age'].mean(), auto_avg(adult['Age'], 1.0)) for _ in range(20)])
r2 = np.mean([pct_error(adult['Capital Gain'].mean(), auto_avg(adult['Capital Gain'], 1.0)) for _ in range(20)])
r3 = np.mean([pct_error(adult['fnlwgt'].mean(), auto_avg(adult['fnlwgt'], 1.0)) for _ in range(20)])

print('Average errors:', r1, r2, r3)
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 1
assert r2 < 100
assert r3 < 1
```

Output:
```
Average errors: 0.020790834901349066 0.007160591377987005 0.009115161105811936
```

## Question 6

In 2-5 sentences each, answer the following:

- Explain your strategy for implementing `auto_avg`
- Argue informally that your definition of `auto_avg` has a total privacy cost of `epsilon`
- Did your solution work well for all three example columns? If it did not work well on any of them, why not?
- When is your solution likely to *not* work well? (i.e. what properties does the data have to have, in order for your solution to not work well?)

<!-- grade_id: cell-0ff63739ed66adf7, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE