# CS 3110/5110: Data Privacy
## Homework 1

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
adult_pii = adult[['Name', 'DOB', 'SSN', 'Zip', 'Age']]
adult_deid = adult.drop(columns=['Name', 'SSN'])
```

## Question 1 (20 points)

Using the dataframes `adult_pii` and `adult_deid`, write code to conduct a linking attack to recover the names of as many individuals in `adult_deid` as possible. Your solution should be parameterized by the set of columns to use in the attack.

<!-- grade_id: cell-e6389ceace711c32, points: 0, grade: false, solution: true -->
```python
def linking_attack(cols):
    return pd.merge(adult_pii, adult_deid, on=cols, how='inner')

```

<!-- grade_id: cell-251b9f1c6f7b4630, points: 20, grade: true, solution: false -->
```python
# TEST CASES for Question 1

assert len(linking_attack(['Zip'])) == 43191
assert len(linking_attack(['Zip', 'DOB'])) == 32563
assert len(linking_attack(['Zip', 'Age'])) == 32755
```

## Question 2 (10 points)

How many individuals in this dataset are uniquely identified by their Zip code? How many are uniquely identified by their age?

Hint: note that the number of *unique ZIP codes* is **different** from the number of *individuals uniquely identified by ZIP code*.

Hint: you can use the `value_counts` method to count the number of occurences of each value in a series.

<!-- grade_id: cell-0e380ac68dd0c77e, points: 0, grade: false, solution: true -->
```python
def unique_zipcode():
    # YOUR CODE HERE
    counts = adult_deid['Zip'].value_counts()
    unique = counts[counts == 1]
    num_unqiue = len(unique)
    return num_unqiue

def unique_dob():
    # YOUR CODE HERE
    counts = adult_deid['DOB'].value_counts()
    unique = counts[counts == 1]
    num_unqiue = len(unique)
    return num_unqiue
```

<!-- grade_id: cell-9b009427c61dbb2b, points: 10, grade: true, solution: false -->
```python
# TEST CASES (INCOMPLETE) for Question 2
# INCOMPLETE means that additional test cases will be used to check your answer during grading

assert unique_zipcode() > 0
assert unique_dob() > 0

```

Output:
```
Zip
64152    1
61523    1
25503    1
75387    1
9820     1
        ..
52272    1
66892    1
94735    1
8213     1
86153    1
Name: count, Length: 23513, dtype: int64
23513
Cell In[145], line 4
      1 # TEST CASES (INCOMPLETE) for Question 2
      2 # INCOMPLETE means that additional test cases will be used to check your answer during grading
      3 
----> 4 assert unique_zipcode() > 0
      5 assert unique_dob() > 0

File ~\Documents\CS3110\cs3110-data-privacy\.venv\Lib\site-packages\pandas\core\generic.py:1501, in NDFrame.__bool__(self)
   1499     @final
   1500     def __bool__(self) -> NoReturn:
-> 1501         raise ValueError(
   1502             f"The truth value of a {type(self).__name__} is ambiguous. "
   1503             "Use a.empty, a.bool(), a.item(), a.any() or a.all()."
   1504         )

ValueError: The truth value of a Series is ambiguous. Use a.empty, a.bool(), a.item(), a.any() or a.all().
```