```python
print('hello')
```

Output:
```
hello
```

# CS 3110/5110: Data Privacy
## In-Class Exercise, week of 8/31/2026

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/refs/heads/main/homework/adult_with_pii.csv')
```

## Question 1

Construct a de-identified version of the `adult` dataset.

<!-- grade_id: cell-4e7a3cdbee680759, points: 0, grade: false, solution: true -->
```python
def deidentify_adult():
    return adult.drop(columns=['Name', 'DOB', 'SSN'])

adult_deid = deidentify_adult()
```

<!-- grade_id: cell-68632b6467ef7923, points: 5, grade: true, solution: false -->
```python
assert 'Name' not in adult_deid.columns
assert 'DOB' not in adult_deid.columns
assert 'SSN' not in adult_deid.columns
```

## Question 2

What columns did you remove from the dataset to de-identify it? Why did you choose these columns? Why did you not choose the remaining columns?

<!-- grade_id: cell-d81a131204234a9d, points: 2, grade: true, solution: true -->
I picked Name, DOB, SSN.

These attributes are certain or nearly certain to uniquely identify individuals in the dataset. Name and SSN definitely do, and DOB might, so let's be on the safe side.

How to decide? Think hard and do your best.

## Question 3

Write a pandas expression to return just the row containing information about Brenn McNeely.

```python
adult[adult['Name'] == 'Brenn McNeely']
# condition: find rows matching a condition: df[? == ?]
# find a row in a dataframe: df[condition]
```

Output:
```
Name       DOB          SSN    Zip  Age Workclass  fnlwgt  \
2  Brenn McNeely  8/6/1991  725-59-9860  95668   38   Private  215646   

  Education  Education-Num Marital Status         Occupation   Relationship  \
2   HS-grad              9       Divorced  Handlers-cleaners  Not-in-family   

    Race   Sex  Capital Gain  Capital Loss  Hours per week        Country  \
2  White  Male             0             0              40  United-States   

  Target  
2  <=50K
```

<!-- grade_id: cell-4baa91fd73c5802e, points: 0, grade: false, solution: true -->
```python
def get_brenns_row():
    return adult[adult['Name'] == 'Brenn McNeely']

brenns_row = get_brenns_row()[['Name', 'DOB', 'Zip', 'Age']]
```

<!-- grade_id: cell-13b75b0658ee9750, points: 2, grade: true, solution: false -->
```python
assert len(brenns_row) == 1
assert brenns_row['Zip'].iloc[0] == 95668
assert brenns_row['Age'].iloc[0] == 38
```

```python
brenns_row
```

## Question 4

Conduct a linking attack to recover Brenn's data from the `adult_deid` dataset.

<!-- grade_id: cell-99305ea08eafa471, points: 0, grade: false, solution: true -->
```python
def recover_brenns_row():
    brenn_age = 38
    brenn_zip = 95668

    brenn_recovered_row = adult_deid[(adult_deid['Age'] == brenn_age) & \
                                     (adult_deid['Zip'] == brenn_zip)]
    return brenn_recovered_row

brenns_recovered_row = recover_brenns_row()
```

<!-- grade_id: cell-b4432ec45f4faf9c, points: 5, grade: true, solution: false -->
```python
assert len(brenns_recovered_row) == 1
assert brenns_recovered_row['Education'].iloc[0] == 'HS-grad'
```

```python

```