# CS 3110/5110: Data Privacy
## In-Class Exercise, week of 9/14/2026

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
```

```python
df = adult[['Education', 'Marital Status', 'Target']]
df.head(10)
```

## Question 6

Imagine the column `Target` is *not* a quasi-identifier, and we generalize the dataset to achieve $k$-Anonymity for $k=2$ as follows:

- Replace each education level below "HS Grad" with `< HS` and others with `>= HS`
- Replace marital status with `Married` and `Not Married`
- Delete rows as required to achieve $k$-anonymity for $k=2$

For which rows is a homogeneity attack possible, and why?

<!-- grade_id: cell-5932f02f09b5cacb, points: 5, grade: true, solution: true -->
Any row where all values for a sensitive value are the same in a set of records. This is because the known values place the target in a specific group and that group membership reveals some new attribute.


## Question 7

Write code to generalize the `Zip` column in the adult dataset.

<!-- grade_id: cell-be9a4264a0e58938, points: 0, grade: false, solution: true -->
```python
def generalize_zip(zip, l):
    return int(str(zip)[:-l] + '0' * l) if l else zip
```

<!-- grade_id: cell-256688fe78b8f7bf, points: 5, grade: true, solution: false -->
```python
# Test cases
assert generalize_zip(47401, 0) == 47401
assert generalize_zip(47401, 1) == 47400
assert generalize_zip(47401, 2) == 47400
assert generalize_zip(47401, 3) == 47000
assert generalize_zip(47401, 4) == 40000
assert generalize_zip(47401, 5) == 0
```

## Question 1

Write a *counting query* to determine whether or not Karrie Trusslove's age is 39.

<!-- grade_id: cell-e9c0c06eb7735ea3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def karrie_query():
    return len(adult[(adult['Name'] == 'Karrie Trusslove') & (adult['Age'] == 39)])
```

<!-- grade_id: cell-f93cfd0eb59130ed, points: 5, grade: true, solution: false -->
```python
# TEST CASE 

assert karrie_query() == 1
```

## Question 2

Add Laplace noise to the counting query you wrote in the last question to ensure differential privacy for $\epsilon = 1.0$.

<!-- grade_id: cell-8ad72c55497bc534, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Obscures whether or not the answer is 0 or 1 by adding random noise
# The problem is that we return the wrong answer
def dp_karrie_query():
    orignal_result = karrie_query()
    noise = np.random.laplace(loc=1, scale=1)
    return orignal_result + noise


```

<!-- grade_id: cell-a7c9bd084e5c5bdf, points: 5, grade: true, solution: false -->
```python
# TEST CASE 

q2_runs = [dp_karrie_query() for _ in range(100)]
noise_runs = [np.random.laplace(loc=1, scale=1/epsilon) for _ in range(100)]

assert stats.wasserstein_distance(q2_runs, noise_runs) < 1
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[7], line 4
      1 # TEST CASE
      2 
      3 q2_runs = [dp_karrie_query() for _ in range(100)]
----> 4 noise_runs = [np.random.laplace(loc=1, scale=1/epsilon) for _ in range(100)]
      5 
      6 assert stats.wasserstein_distance(q2_runs, noise_runs) < 1

NameError: name 'epsilon' is not defined
```

## Question 3

In 2-4 sentences, describe how `dp_karrie_query` protects Karrie's privacy.

<!-- grade_id: cell-07337a085c0ba0e1, points: 5, grade: true, solution: true -->
dp_karrie_query protects Karrie's privacy by adding random noise to here data. This works by 

## Question 4

Implement a function to construct an *empirical distribution* over random samples from the Laplace distribution with location 0 and scale `scale`. Use 1000 samples.

<!-- grade_id: cell-e7ddf0e4b195dfdc, points: 0, grade: false, solution: true -->
```python
def empirical_dist_lap(scale):
    # YOUR CODE HERE
    raise NotImplementedError()

epsilon1 = 1.0
epsilon2 = 2.0

plt.hist(empirical_dist_lap(1/epsilon1), bins=50);
plt.hist(empirical_dist_lap(1/epsilon2), bins=50, alpha=.7);
```

<!-- grade_id: cell-e68891bb70939fd0, points: 5, grade: true, solution: false -->
```python
# PLACEHOLDER for test cases

```

## Question 5

Consider the program below.

```python
counting_query = lambda: 1

result_1 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon1)
result_2 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon2)
result_3 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon1)

print(result_1, result_2, result_3)
```

Write a function `total_epsilon` that returns the total privacy cost for this program.

<!-- grade_id: cell-32b1cdb0a2990e8f, points: 0, grade: false, solution: true -->
```python
def total_epsilon():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e484920fc1429ad0, points: 5, grade: true, solution: false -->
```python
# Placeholder for test case

```